# SLU 17_2 - Exam Prep I

## Batch 2 Python Exam 

This is the Python admission exam from the Lisbon Data Science Starters Academy - Batch 2. The allocated time for this exam was 2 hours. If you plan to take the admission exam this year, it's a good idea to measure the time you need to solve this notebook. You can see how far you get in 2 hours and also how much time you need to complete the notebook. At the same time, it's a good idea to not rush through the exercises and rather concentrate on working thoroughly.

At the end, you submit the notebook to the portal as usual. You can submit as many times as you like.

```python
from operator import itemgetter

#used for evaluation
import inspect
import math
import re
import json
import sys
import utils
import random

#generate data
data = utils.generate_data()
data_comma =  utils.generate_data_comma()
data_clean = utils.generate_data_clean()
```

## Question 1

If you took a look at the dataset, you may have noticed that some of the numeric values are "polluted" with characters.

In this task you'll have to write a function, `clean_percentage`, that should do the following:

- Receive as argument a list of dictionaries;
 - Each of the dictionaries keeps the % of female students in higher education per area, for a specific year;
- Remove all the characters and spaces from the numeric values;
- Return the cleaned data;

Hint: mind the naughty (R)'s :)

<!-- grade_id: cell-c7177b04bd669cff, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def clean_percentage(data):
    for i in range(0, len(data)):
        for key in data[i].keys():
            if type(data[i][key]) is str and data[i][key].isdigit() == False:
                lst = data[i][key].split()
                for word in lst:
                    if all(char in "0123456789,." for char in word):
                        data[i][key] = word
    return data
```

<!-- grade_id: cell-18bc7007043d3ec2, points: 3, grade: true, solution: false -->
```python
#regenerate data
data = utils.generate_data()
data_comma =  utils.generate_data_comma()
data_clean = utils.generate_data_clean()

data_fixed = clean_percentage(data)
assert len(data_fixed) == 10, "Not correct. Keep trying."
assert isinstance(data_fixed, list), "Not correct. Keep trying."

for r in range(len(data)):
    assert isinstance(data[r], dict), "Not correct. Keep trying."
    assert len(data[r]) == 10, "Not correct. Keep trying."

    for k in data[r].keys():
        assert "(R)" not in data_fixed[r][k], "Not correct. Keep trying."
        assert re.search(r"[a-zA-Z\(\)]", data_fixed[r][k]) is None, "Not correct. Keep trying."
        
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 2

In this task, we want to clean the names of the education areas, so they are all lower case and without white spaces.

For instance, Social Sciences, Management and Law should become social_sciences_management_and_law.

So you'll have to write two functions.

The first function, `clean_header_string`, should do the following:

- Receive a string as input;
- Remove all the commas from the string;
- Replace the white spaces in the string with underscores;
- Convert the string to lower case;
- Return the transformed string;

The second function, `clean_header`, should do the following:

- Receive as input, a list of dictionaries, similar to the one in the last question;
- Use the first function to transform the education area names;
- Return the transformed data;


<!-- grade_id: cell-a11732fd19b01c68, points: 0, grade: false, solution: true -->
```python
def clean_header_string(string):
    # YOUR CODE HERE
    res = ""
    for char in string:
        if char == ",":
            continue
        else:
            res += char
    res2 = ""
    for char in res:
        if char == " ":
            res2 += "_"
        else:
            res2 += char
    res2 = res2.lower()
    return res2

def clean_header(d):
    # YOUR CODE HERE
    lst = []
    for i in range(0, len(data)):
        new_dict = {clean_header_string(key):value for key, value in data[i].items()}
        lst.append(new_dict)
    return lst

data
```

Output:
```
[{'Year': '2013',
  'Total': '53,2',
  'Education': '80,4',
  'Arts and Humanities': '56,9',
  'Social Sciences, Management and Law': '58,1',
  'Sciences, Maths and Computer Science': '47,2',
  'Engineering': '27,4',
  'Agriculture': '56,9',
  'Health and Social Security': '76,8',
  'Services': '41,4'},
 {'Year': '201',
  'Total': '22,2',
  'Education': '80,4',
  'Arts and Humanities': '16,9',
  'Social Sciences, Management and Law': '58,1',
  'Sciences, Maths and Computer Science': '47,2',
  'Engineering': '27,4',
  'Agriculture': '56,9',
  'Health and Social Security': '76,8',
  'Services': '41,4'},
 {'Year': '2014',
  'Total': '53,5',
  'Education': '80,7',
  'Arts and Humanities': '58,0',
  'Social Sciences, Management and Law': '58,5',
  'Sciences, Maths and Computer Science': '47,5',
  'Engineering': '27,6',
  'Agriculture': '56,9',
  'Health and Social Security': '76,6',
  'Services': '41,1'},
 {'Year': '2015',
  'Total': '53,6',
  'Education': '80,7',
  'Arts and Humanities': '
```

<!-- grade_id: cell-d7215a1617be8293, points: 3, grade: true, solution: false -->
```python
#regenerate data
data = utils.generate_data()
data_comma =  utils.generate_data_comma()
data_clean = utils.generate_data_clean()

for r in range(len(data)):
    assert len(data) == 10, "Not correct. Keep trying."
    assert len(data[r]) == 10, "Not correct. Keep trying."
    for k in data[r].keys():
        area_clean = clean_header_string(k)
        assert isinstance(area_clean, str), "Not correct. Keep trying."
        assert re.search(r"[\,\sA-Z]", area_clean) is None, \
               "Not correct. Keep trying."

data_fixed = clean_header(data)
assert isinstance(data_fixed, list), "Not correct. Keep trying."
for r in range(len(data_fixed)):
    assert len(data_fixed) == 10, "Not correct. Keep trying."
    assert len(data_fixed[r]) == 10, "Not correct. Keep trying."
    assert isinstance(data_fixed[r], dict), "Not correct. Keep trying."
    for k in data_fixed[r].keys():
        assert isinstance(k, str), "Not correct. Keep trying."
        assert re.search(r"[\,\sA-Z]", k) is None, \
               "Not correct. Keep trying."
        
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 3

In this task, we want to convert our data to the right data types.

For this you'll have to write two functions.


The first function, `commas`, should do the following:

- Receive as input a list of dictionaries, similar to the previous questions;
- Replace the commas with dots in the numeric values;
- Return the transformed data;


The second function `data_types` should do the following:

- Receive as input a list of dictionaries (specifically the one returned by the first function);
- Convert the year values to integer;
- Convert the % of female students values to float;
- Return the transformed data;

<!-- grade_id: cell-76ca0448c5209f19, points: 0, grade: false, solution: true -->
```python
def commas(d):
    # YOUR CODE HERE
    for i in range(0, len(d)):
        for key in d[i].keys():
            d[i][key] = d[i][key].replace(",", ".")
    return d

def data_types(d):
    # YOUR CODE HERE
    d = clean_header(d)
    d = commas(d)
    d = clean_percentage(d)
    for i in range(0, len(d)):
        for key in d[i].keys():
            if key == "year":
                d[i][key] = int(d[i][key])
            else:
                d[i][key] = float(d[i][key])
    return d

data
```

Output:
```
[{'Year': '2013',
  'Total': '53,2',
  'Education': '80,4',
  'Arts and Humanities': '56,9',
  'Social Sciences, Management and Law': '58,1',
  'Sciences, Maths and Computer Science': '(R) 47,2',
  'Engineering': '(R) 27,4',
  'Agriculture': '56,9',
  'Health and Social Security': '76,8',
  'Services': '(R) 41,4'},
 {'Year': '201',
  'Total': '22,2',
  'Education': '80,4',
  'Arts and Humanities': '16,9',
  'Social Sciences, Management and Law': '58,1',
  'Sciences, Maths and Computer Science': '(R) 47,2',
  'Engineering': '(R) 27,4',
  'Agriculture': '56,9',
  'Health and Social Security': '76,8',
  'Services': '(R) 41,4'},
 {'Year': '2014',
  'Total': '53,5',
  'Education': '80,7',
  'Arts and Humanities': '58,0',
  'Social Sciences, Management and Law': '58,5',
  'Sciences, Maths and Computer Science': '(R) 47,5',
  'Engineering': '(R) 27,6',
  'Agriculture': '56,9',
  'Health and Social Security': '76,6',
  'Services': '(R) 41,1'},
 {'Year': '2015',
  'Total': '53,6',
  'Education'
```

<!-- grade_id: cell-c0f514108d4869b2, points: 3, grade: true, solution: false -->
```python
#regenerate data
data = utils.generate_data()
data_comma =  utils.generate_data_comma()
data_clean = utils.generate_data_clean()

data_dot = commas(data_comma)
assert isinstance(data_dot, list), "Not correct. Keep trying."
for r in range(len(data_dot)):
    assert len(data_dot) == 10, "Not correct. Keep trying."
    assert len(data_dot[r]) == 10, "Not correct. Keep trying."
    assert isinstance(data_dot[r], dict), "Not correct. Keep trying."
    for v in data_dot[r].values():
        assert isinstance(v, str), "Not correct. Keep trying."
        assert re.search(r"[\,]", v) is None, \
               "Not correct. Keep trying."

data_dot_type = data_types(data_dot)
assert isinstance(data_dot_type, list), "Not correct. Keep trying."
for r in range(len(data_dot_type)):
    assert len(data_dot_type) == 10, "Not correct. Keep trying."
    assert len(data_dot_type[r]) == 10, "Not correct. Keep trying."
    assert isinstance(data_dot_type[r], dict), "Not correct. Keep trying."
    for k, v in data_dot_type[r].items():
        if k == "year":
            assert isinstance(v, int), "Not correct. Keep trying."
        else:
            assert isinstance(v, float), "Not correct. Keep trying."
            
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 4

In this task, we want to focus on data about a specific education area.

So you'll have to write a function, `education_years`  that should do the following:

- Receive as input a list of dictionaries (similar to the ones of the previous questions), and an education area (string);
- Return a list of tuples where the first value of the tuple is the year and the second value is the % of females in the given education area;

In this task you're required to use the following:

- `map` function;
- `lambda` function;

<!-- grade_id: cell-0546a2b6445fd75b, points: 0, grade: false, solution: true -->
```python
def education_years(data, education_area):
    # YOUR CODE HERE
    lst = list(map(lambda key: (key["year"], key[education_area]), data))
    return lst

```

<!-- grade_id: cell-0bf157ce02b1b085, points: 3, grade: true, solution: false -->
```python
#regenerate data
data = utils.generate_data()
data_comma =  utils.generate_data_comma()
data_clean = utils.generate_data_clean()

education_areas = [key for key in data_clean[0].keys() if key not in ["year"]]
for area in education_areas:
    percent_women_year = education_years(data_clean, area)
    assert isinstance(percent_women_year, list), "Not correct. Keep trying."
    assert len(percent_women_year) == 12, "Not correct. Keep trying."
    for year in percent_women_year:
        assert isinstance(year, tuple), "Not correct. Keep trying."
        assert len(year) == 2, "Not correct. Keep trying."
        assert isinstance(year[0], int), "Not correct. Keep trying."
        assert isinstance(year[1], float), "Not correct. Keep trying."

test1 = education_years(data_clean, "total")
assert math.isclose(test1[3][1], 53.5, abs_tol=0.1), "Not correct. Keep trying."
assert test1[10][0] == 2016, "Not correct. Keep trying."

test2 = education_years(data_clean, "services")
assert math.isclose(test2[5][1], 41.7, abs_tol=0.1), "Not correct. Keep trying."
assert test2[8][0] == 2014, "Not correct. Keep trying."

source = inspect.getsource(education_years)
assert "map" in source, "Not correct. Keep trying."
assert "lambda" in source, "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 5

In this task, for a certain education area, we want to understand what were the years where more females were enrolled in higher education.

So you'll have to write a function named `female_enrolled` that should do the following:

- Receive as input a list of tuples with year and % of females (like the output of last task) and a threshold;
- Filter **out** all the tuples that correspond to a % of females strictly below the threshold;
- Sort the remaining tuples from higher to lower % of females;
- Return the resulting list of tuples;

In this task you're required to use the following:

- `sorted` function

Hint: `itemgetter` may be handy

<!-- grade_id: cell-3c339e619872e065, points: 0, grade: false, solution: true -->
```python
import operator 

def female_enrolled(data, threshold):
    # YOUR CODE HERE
    lst = []
    for i in range(0, len(data)):
        if data[i][1] >= threshold:
            lst.append(data[i])
    lst = sorted(lst, key=operator.itemgetter(1), reverse=True)
    return  lst

dta = [(2100, 80), (2010, 70), (2030, 90)]
dta = female_enrolled(dta, 80)
dta
```

Output:
```
[(2030, 90), (2100, 80)]
```

<!-- grade_id: cell-30f67a7608d49607, points: 3, grade: true, solution: false -->
```python
#regenerate data
data = utils.generate_data()
data_comma =  utils.generate_data_comma()
data_clean = utils.generate_data_clean()

data_enrolled = utils.generate_data_enrolled()
data_filtered = female_enrolled(data_enrolled, 0)
assert isinstance(data_filtered, list), "Not correct. Keep trying."
assert len(data_filtered) == 5, "Not correct. Keep trying."
for year in data_filtered:
    assert isinstance(year, tuple), "Not correct. Keep trying."
    assert len(year) == 2, "Not correct. Keep trying."
    assert isinstance(year[0], int), "Not correct. Keep trying."
    assert isinstance(year[1], float), "Not correct. Keep trying."

test1 = female_enrolled(data_enrolled, 58.1)
assert len(test1) == 3, "Not correct. Keep trying."
assert math.isclose(test1[1][1],
                    58.7,
                    abs_tol=0.1), "Not correct. Keep trying."
assert test1[1][0] == 2015, "Not correct. Keep trying."

test2 = female_enrolled(data_enrolled, 58.0)
assert len(test2) == 4, "Not correct. Keep trying."
assert math.isclose(test2[3][1],
                    58.0,
                    abs_tol=0.1), "Not correct. Keep trying."
assert test2[3][0] == 2014, "Not correct. Keep trying."

source = inspect.getsource(female_enrolled)
assert "sorted" in source, "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 6

In this task we'll create a class to play the lottery!

First, get to know the `raffle_key`. The `raffle_key` is a dictionary, associating ticket numbers with prizes. All the numbers that are not in the raffle_key won't get any prize.

```python
raffle_key = {1:"10 euros", 5:"100 euros", 9:"40 euros"}
```

### Step 1:

- Create a `Lottery` class that takes one constructor parameter called `raffle_key`;
- In the `Lottery` class create a method named `get_lucky` that receives the lucky_number as parameter;
    and returns the strings "{X} euros" or "Better luck next time!", depending on the presence of the lucky number in the raffle key;
- In our example above, our lucky number 9 corresponds to the response string "40 euros";
    if in some case the lucky number provided doesn't exist in the raffle key, we should print out "Better luck next time!";
- If no lucky number is provided whatsoever, make sure that a random one is generated.


### Step 2:

- Instantiate the `Lottery` class with name `simple_lottery` and with the given raffle key;


In this task you're required to use the following:

- `try/except` statements;


<!-- grade_id: cell-dc8077515454abe0, points: 0, grade: false, solution: true -->
```python
class Lottery():
    # YOUR CODE HERE
    def __init__(self, raffle_key):
        self.raffle_key = raffle_key
    
    def get_lucky(self, lucky_number=None):
        if lucky_number is None:
            lucky_number = random.randint(1, 100)
        try:
            raffle_key[lucky_number]
        except:
            return "Better luck next time!"
        return raffle_key[lucky_number]
    
```

Output:
```
'40 euros'
```

<!-- grade_id: cell-96c7d9be71562f8a, points: 5, grade: true, solution: false -->
```python
assert isinstance(simple_lottery, Lottery), "Not correct. Keep trying."
assert hasattr(simple_lottery, "raffle_key"), "Not correct. Keep trying."
assert simple_lottery.get_lucky(1) == "10 euros", "Not correct. Keep trying."
assert simple_lottery.get_lucky(5) == "100 euros", "Not correct. Keep trying."
assert simple_lottery.get_lucky(9) == "40 euros", "Not correct. Keep trying."
assert simple_lottery.get_lucky("1") == "Better luck next time!", \
       "Not correct. Keep trying."
assert simple_lottery.get_lucky(-3) == "Better luck next time!", \
       "Not correct. Keep trying."
assert isinstance(simple_lottery.get_lucky(), str), "Not correct. Keep trying."
source = inspect.getsource(simple_lottery.get_lucky)
assert "try" in source, "Not correct. Keep trying."
assert "except" in source, "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

### Submit your work!

To submit your work, [follow the steps here, in the step \"Grading the Exercise Notebook\"!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)